![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

# Teste 1: Bisseção

```python
import numpy as np
```

## Questão 1: Número de iterações da bisseção, e de chamadas à função `f`

Generalize o algoritmo da bisseção para retornar, nesta ordem:
- a (aproximação da) raiz,
- o número de bisseções (divisões por 2 do intervalo) feitas,
- o número de vezes que você chamou a função `f`.

Use como critério de parada a tolerância `tol`,
ou seja, o algoritmo termina quando:
- seja possível garantir que o erro (absoluto) da resposta ("em $x$") seja menor do que `tol`.

<!-- grade_id: cell-bfd2f7379cc308c4, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, tol=1e-8):
    va = f(a)
    if(va == 0):
        return (a, 0, 1)
    else:
        vb = f(b)
        if(vb == 0):
            return (b, 0, 2)
        else:
            top = a
            bottom = b
            ncalls = 2
            niters = 0
            while np.abs(top - bottom) > tol:
                i = (top + bottom) / 2
                vi = f(i)
                if vi * va < 0:
                    bottom = i
                else:
                    top = i
                ncalls += 1
                niters += 1
            return (i, niters, ncalls)
```

### Testes simples

```python
def p2(x): return x**2 - 2
```

<!-- grade_id: cell-c62e828e23ba0c73, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2)
assert abs(raiz - np.sqrt(2)) < 1e-8
assert 26 <= niters <= 30
```

<!-- grade_id: cell-3b33b89bcafe1954, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2, tol=1e-3)
assert abs(raiz - np.sqrt(2)) < 1e-3
assert 10 <= niters <= 15
```

<!-- grade_id: cell-432c2ad26a7a7d7e, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2, tol=1e-3)
assert abs(raiz - np.sqrt(2)) > 3e-4
```

<!-- grade_id: cell-721772ed44e72acb, points: 2, grade: true, solution: false -->
```python
# Testes escondidos aqui
```

### Testes mais legais

<!-- grade_id: cell-8d7ce0c5727c5838, points: 1, grade: true, solution: false -->
```python
pimeios, niters, ncalls = bissecao(np.cos, 0, 2)
assert abs(pimeios - np.pi/2) <= 1e-8
```

<!-- grade_id: cell-ff05771e68d86f63, points: 1, grade: true, solution: false -->
```python
menospimeios, niters, ncalls = bissecao(np.cos, -2, 0)
assert abs(menospimeios + np.pi/2) <= 1e-8
```

## Questão 2: várias iterações da bisseção

```python
print(f'{"raiz":^9} {"niters"} {"ncalls"}, {"erro":^11} {"tol":^10}')
for p in range(10,20):
    raiz, niters, ncalls = bissecao(p2, 0, 2, tol=2**(-p))
    print(f"{raiz:.7f} {niters:6d} {ncalls:6d}, {raiz-np.sqrt(2): .8f} {2**(-p):.8f}")
```

Output:
```
raiz    niters ncalls,    erro        tol    
1.4150391     11     13,  0.00082550 0.00097656
1.4145508     12     14,  0.00033722 0.00048828
1.4143066     13     15,  0.00009308 0.00024414
1.4141846     14     16, -0.00002899 0.00012207
1.4142456     15     17,  0.00003204 0.00006104
1.4142151     16     18,  0.00000153 0.00003052
1.4141998     17     19, -0.00001373 0.00001526
1.4142075     18     20, -0.00000610 0.00000763
1.4142113     19     21, -0.00000229 0.00000381
1.4142132     20     22, -0.00000038 0.00000191
```

O que você observa na tabela acima?  Explique.

<!-- grade_id: cell-2f026a9298f20d3b, points: 3, grade: true, solution: true -->
É possível ver, acima, que enquanto o valor de tolerância reduz-se exponencialmente (pois trata-se de uma potência inversa de dois), o valor de erro também reduz-se, aproximadamente, de forma exponencial inversa, enquanto o número de iterações e de chamadas de função aumenta linearmente, para se chegar ao nível de tolerância e erro esperado.

Esse comportamento ocorre, porque o processo de bisseção é exponencial inverso, sempre reduzindo o intervalo de avaliação pela metade, em um processo exponencial, que, em relação à variação exponencial da tolerância, segundo a tabela acima, é linear (pois a variação de tolerância dada acima é linear).

Se a variação da tolerância fosse linear, na tabela acima, a variação de iterações e chamadas seria logarítmica.


## Questão 3: Outros critérios de parada

Modifique a sua bisseção para ter 2 critérios de parada:
- quando o tamanho do intervalo ficar menor do que `tol`
- quando o valor da função ficar menor do que `err`

<!-- grade_id: cell-89ea2904b3c908eb, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, tol=1e-8, err=1e-8):
    va = f(a)
    if(va == 0):
        return (a, 0, 1)
    else:
        vb = f(b)
        if(vb == 0):
            return (b, 0, 2)
        else:
            vi = float('inf')
            top = a
            bottom = b
            ncalls = 2
            niters = 0
            while np.abs(top - bottom) > tol and np.abs(vi) > err:
                i = (top + bottom) / 2
                vi = f(i)
                if vi * va < 0:
                    bottom = i
                else:
                    top = i
                ncalls += 1
                niters += 1
            return (i, niters, ncalls)
```

Pare e pense:

uma solução iterativa pode ficar mais fácil de adaptar para vários critérios de parada
se ela for escrita com `while True`.

### Testes

<!-- grade_id: cell-e5dbbdc7bb04bd36, points: 2, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2)
assert abs(p2(raiz) - 2) > 1e-8
assert abs(raiz - np.sqrt(2)) < 1e-8
```

<!-- grade_id: cell-65495ebe135fbf41, points: 1, grade: true, solution: false -->
```python
assert bissecao(np.cos, 0, 2) != bissecao(np.cos, 0, 2, err=0)
```

<!-- grade_id: cell-9446992c617f02eb, points: 1, grade: true, solution: false -->
```python
assert bissecao(np.cos, 0, 2) == bissecao(np.cos, 0, 2, tol=0)
```

O que os dois testes acima indicam sobre a bisseção para `np.cos`?

<!-- grade_id: cell-d8ce30660516f95d, points: 2, grade: true, solution: true -->
O processo de bissecção, quando se utiliza o critério da tolerância, permite à chegada de um intervalo mínimo, no qual se esgotam todas as possibilidades de continuar a bissecção, porque atingiu-se o valor máximo da precisão da representação do número flutuante, o que não é necessariamente verdade com o critério erro. Isso também significa que o critério de tolerância converge mais rapidamente do que o critério de erro.

## Questão 4: Bisseção sem função

Escreva uma função que divide o intervalo ao meio, fica com o esquerdo,
e repete até o limite de precisão do computador.

Ela deve retornar o último intervalo não-trivial (como uma tupla), e o número de divisões.

<!-- grade_id: cell-dbe61010623d813f, points: 0, grade: false, solution: true -->
```python
def bisseção_esquerda(a,b):
    if a == b:
        return (a, b, 0)
    if a < b:
        left = a
        right = b
    else:
        left = b
        right = a
    niters = 0
    prev_right = None
    while np.abs(right - left) > 0 and right != prev_right:
        prev_right = right
        right = (right + left) / 2 
        niters += 1
    return ((left, prev_right), niters-1)
```

<!-- grade_id: cell-365cb4391a9f65a4, points: 1, grade: true, solution: false -->
```python
(an, bn), ndivs = bisseção_esquerda(1,2)
assert 2e-16 <= bn - an <= 4e-16
assert ndivs == 52
```

<!-- grade_id: cell-179c89a09dfda258, points: 1, grade: true, solution: false -->
```python
(an, bn), ndivs = bisseção_esquerda(bn,2)
assert 2e-16 <= bn - an <= 4e-16
assert ndivs == 51
```

Use esta função para calcular o próximo número de ponto flutuante.

<!-- grade_id: cell-ace3480d37572adb, points: 0, grade: false, solution: true -->
```python
def próximo_float(x):
    (an, bn), ndivs = bisseção_esquerda(x,x + np.abs(x) + 1)
    return bn
```

```python
próximo_float(1)
```

Output:
```
1.0000000000000002
```

<!-- grade_id: cell-8f75b70a7dbab81a, points: 1, grade: true, solution: false -->
```python
nextone = próximo_float(1)
assert nextone > 1
assert (1+nextone)/2 == 1
```

<!-- grade_id: cell-b39b0becc0922ad6, points: 1, grade: true, solution: false -->
```python
nextzero = próximo_float(0)
assert nextzero > 0
assert nextzero/2 == 0
```

<!-- grade_id: cell-c902bc7ab832c0d0, points: 1, grade: true, solution: false -->
```python
nextnext = próximo_float(próximo_float(1))
assert nextnext > 1
assert 1 < (nextnext+1)/2 < nextnext
```

<!-- grade_id: cell-96bd65cd68e38dc0, points: 1, grade: true, solution: false -->
```python
maisum = próximo_float(-1.435)
assert maisum > -1.435
assert (maisum - 1.435)/2 == -1.435
```

## Questão 5: bisseção alternante

Escreva uma função que escolhe o lado esquerdo, e depois o lado direito, e repete, até acabar.

<!-- grade_id: cell-3a6605b0c2fff49d, points: 0, grade: false, solution: true -->
```python
def bisseção_alt(a,b):
    if a == b:
        return (a, b, 0)
    if a < b:
        left = a
        right = b
    else:
        left = b
        right = a
    niters = 0
    to_left = True
    prev_left = None
    prev_right = None
    while np.abs(right - left) > 0 and ((to_left and left != prev_left) or ((not to_left) and right != prev_right)):    
        prev_right = right
        prev_left = left
        if to_left:
            right = (right + left) / 2
        else:
            left = (right + left) / 2
        to_left = not to_left
        niters += 1
    return ((prev_left, prev_right), niters-1)
```

<!-- grade_id: cell-d1906fd95ebb85a3, points: 1, grade: true, solution: false -->
```python
(a,b), ndivs = bisseção_alt(1,2)
assert 1 < a < b < 2
assert ndivs == 52
```

<!-- grade_id: cell-1433725a5fa659f7, points: 1, grade: true, solution: false -->
```python
(a,b), ndivs = bisseção_alt(1,4)
assert 1 < a < b < 4
assert ndivs == 53
assert (a+b)/2 == b
```

```python
for (a0,b0) in [(1,2), (1,4), (2,4), (1/3,3/5)]:
    (a,b), ndivs = bisseção_alt(a0,b0)
    print(f"{a0:.5f} {b0:.5f} ---> {a:.17f} {b:.17f}")
```

Output:
```
1.00000 2.00000 ---> 1.33333333333333326 1.33333333333333348
1.00000 4.00000 ---> 1.99999999999999978 2.00000000000000000
2.00000 4.00000 ---> 2.66666666666666652 2.66666666666666696
0.33333 0.60000 ---> 0.42222222222222267 0.42222222222222272
```

O que esta função parece calcular?  Explique.

<!-- grade_id: cell-4d30f5821b51b3d7, points: 3, grade: true, solution: true -->
A função calcula o valor que está a um terço do ponto à esquerda, considerando o intervalo (a, b) de entrada da função. Isso ocorre, porque a alternância dos intervalos leva à seguinte sequência numérica, que converge para 1/3.

Vamos considerar a distância entre a e b definida por d (parâmetros da função de bisseção alternada):

$$ d = b - a $$

Claramente, o valor de b é dado por:

$$ b = a + d $$

Já o valor do ponto encontrado pela bisseção alternada pode ser dada pela sequência alternada de distâncias:

$ d_f = d - \frac{1}{2}d - \frac{1}{4}d + \frac{1}{8}d - \frac{1}{16}d + \frac{1}{32}d - \frac{1}{64}d + ... = $

$ d(1 - \frac{1}{2} - \frac{1}{4} + \frac{1}{8} - \frac{1}{16} + \frac{1}{32} - \frac{1}{64} + ...) = $

$ d(\frac{1}{2} - \frac{1}{4} + \frac{1}{8} - \frac{1}{16} + \frac{1}{32} - \frac{1}{64} + ...) = $

$ d(\frac{\frac{1}{2}}{1 - (-\frac{1}{2})}) = \frac{1}{3}d $

O que implica qhe o ponto final da sequência é dado por $b = a + \frac{1}{3}d = a + \frac{b-a}{3}$

<!-- grade_id: cell-36e18b27847e0f36, points: 1, grade: true, solution: false -->
```python
(a,b), ndivs = bisseção_alt(-1,2)
assert -1 < a < b < 2
assert (a+b)/2 == b
assert ndivs > 1000
```

Esta iteração levou mais de 1000 bisseções.  Explique o que aconteceu.

<!-- grade_id: cell-203e11f457638867, points: 2, grade: true, solution: true -->
O intervalo entre $ [-1, 1] $, na representação de números de ponto flutuante, é o intervalo que possui a maior quantidade possível de números com representação válida. Na verdade, o intervalo $ [-1, 1] $ possui tantas representações de números flutuantes quanto todos os outros intervalos de números da reta real unidos, levando a bisseção ser muito mais profunda do que em outros intervalos.


```python

```